---
title       : Insert the chapter title here
description : Insert the chapter description here
--- type:ConsoleExercise lang:python xp:100 skills:2 key:9992d4d90e
## <<<New Exercise>>>


*** =instructions

*** =hint

*** =pre_exercise_code
```{python}

```

*** =sample_code
```{python}

```

*** =solution
```{python}

```
print ("hello")

*** =sct
```{python}

```
